# Lab 1

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab_1.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

## Lab 1: Functions


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Fix the Bug

<hr style="border: 1px solid #fdb515;" />

The following snippet of code doesn't work! Figure out what is wrong and fix the bugs.

```python
def both_positive(x, y):
    # Returns True if both x and y are positive.
    # >>> both_positive(-1, 1)
    # False
    # >>> both_positive(1, 1)
    # True
    return x and y > 0 # You can replace this line!

```

```python
grader.check("both_positive")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Sum Digits

<hr style="border: 1px solid #fdb515;" />

Write a function that takes in a nonnegative integer and sums its digits. (Using
floor division and modulo might be helpful here!)

**Hint**: Consider how numbers are represented in decimal form, with ones, tens, hundreds, etc. What can we use to floor/mod a value to get the ones/tens/hundreds etc. digit?

```python
def sum_digits(n):
    # Sum all the digits of n.
    # >>> sum_digits(5)
    # 5
    # >>> sum_digits(34) # 3 + 4 = 7
    # 7
    # >>> sum_digits(10) # 1 + 0 = 1
    # 1
    # >>> sum_digits(4224) # 4 + 2 + 2 + 4 = 12
    # 12
    # >>> sum_digits(1234567890)
    # 45
    # >>> x = sum_digits(123) # make sure that you are using return rather than print
    # >>> x
    # 6
    "*** YOUR CODE HERE ***"

```

```python
grader.check("sum_digits")
```

Tests for both_positive (points: 1)
```
>>> both_positive(-1, 1) == False
True
```
```
>>> both_positive(1, 1) == True
True
```


Tests for sum_digits (points: 1)
```
>>> sum_digits(5) == 5
True
```
```
>>> sum_digits(34) == 7
True
```
```
>>> sum_digits(10) == 1
True
```
```
>>> sum_digits(4224) == 12
True
```
```
>>> sum_digits(1234567890) == 45
True
```
```
>>> x = sum_digits(123)
>>> x == 6
True
```
